```python
%%capture
!pip install otter-grader
```

```python
import otter
grader = otter.Notebook(colab = True)
```

# Korzystanie z Google Colab i Jupyter Notebooks

Podczas kursu będziemy korzystać z Notebooków taki jak ten. Pozwoli to nam przejść od razu do nauki programowania bez potrzeby instalacji oprogramowania na własnym komputerze.

# Pytania zamknięte

Podczas pracy z notebookiem co jakiś czas napotkasz krótkie zadania do rozwiązania. Na przykład:

**Zadanie 1** Zdefiniuj funkcję `square`, która liczy kwadrat podanej liczby

```python
def square(x):
    # Umieść swój kod poniżej, a potem urochom te komórkę.
    # Następnie wykonaj kod z następnej komórki - Twoje rozwiązanie zostanie automatycznie przetestowane.
    ...
```

```python
grader.check("q1")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: None)
```
>>> square(1)==1
True
```
```
>>> square(0.5)==(0.25)
True
```
```
>>> square(2)==4
True
```
